# Numpy


```python
%pip install -q otter-grader

import numpy as np
import otter

grader = otter.Notebook()

%matplotlib inline
```

**Aufgabe 1:** Erstellen Sie ein Numpy-Array mit den Werten 1, 3, 8 und speichern Sie es in der Variable `a`. Definieren Sie ein zweites Numpy-Array, indem Sie 2 zu jedem Element von `a` addieren und speichern Sie es in der Variable `b`.

```python
a = np.array([1, ..., ...])
b = a + ...
```

```python
grader.check("q1")
```

**Aufgabe 2:** Im Folgenden werden zwei Numpy `x1` und `x2` definiert. Speichen Sie das Quotient von `x1` und `x2` in der Variable `x1_div_x2` ab.

```python

```python
x1 = np.array([1, 3, 8, 6])
x2 = np.array([2, 5, 3, 7])
x1_div_x2 = ... / ...
```

```python
grader.check("q2")
```

**Aufgabe 3:** Erstellen Sie ein Numpy-Array mit den Werten "a", "b", "c" und "d" und speichern Sie es in der Variable `my_array`. Speichern Sie das dritte Element von `my_array` in der Variable `third_element`.

**Hinweis:** Die Indizierung in Python beginnt bei 0. Das erste Element eines Arrays hat also den Index 0.

```python
import numpy as np
arr = np.array([32, 3, 8])
print(arr[0]) # Gibt 32 aus

print(arr[1]) # Gibt 3 aus
```

```python
third_element = ...[...]
third_element
```

```python
grader.check("q3")
```

**Aufgabe 4:** Die Variable `sent` enthält eine einige Sätze aus dem Buch "Alice in Wonderland". Estellen Sie ein Numpy-Array, das die einzelnen Zeichen des Strings `sent` enthält und speichern Sie es in der Variable `sent_arr`.

In Python können Sie über die einzelnen Zeichen eines Strings iterieren, indem Sie den String in einer Schleife durchlaufen. Der folgende Code gibt die einzelnen Zeichen des Strings `sent` aus:

```python
for char in sent:
    print(char)
```

**Hinweis:** Sie können die Funktion `np.array()` verwenden, um ein Numpy-Array zu erstellen.
    
```python
sent = "A place to be"
sent_arr = np.array([char for char in sent])
```

Wählen Sie jedes dritte Element des Arrays `sent_arr` und speichern Sie es in der Variable `every_third_char`. Wie viele verschiedene Zeichen enthält `every_third_char`?
Speichern Sie das Ergebnis in der Variable `num_unique_chars`.

**Hinweis**: Sie können die Funktion `np.unique()` verwenden, um die eindeutigen Elemente eines Numpy-Arrays zu finden.

```python
unique_elements = np.unique(np.array([1, 1, 2]))
print(unique_elements) # Gibt [1, 2] aus
```


```python
sent = """
There was nothing so very remarkable in that; nor did Alice think it so very much 
out of the way to hear the Rabbit say to itself, “Oh dear! Oh dear! I shall be late!” 
(when she thought it over afterwards, it occurred to her that she ought to have wondered at this, but at the time it all seemed quite natural); but when the Rabbit actually took a watch out of its waistcoat-pocket, and looked at it, and then hurried on, Alice started to her feet, for it flashed across her mind that she had never before seen a rabbit with either a waistcoat-pocket, or a watch to take out of it, and burning with curiosity, she ran across the field after it,
and fortunately was just in time to see it pop down a large rabbit-hole under the hedge.
"""
sent_arr = np.array(...)
every_third_char = sent_arr[...]
num_unique_chars = every_third_char...
```

```python
grader.check("q4")
```

**Aufgabe 5:** Erstellen Sie ein Numpy-Array, das 500 mal die Zahl 1 enthält und speichern Sie es in der Variable `my_ones`. Erstellen Sie ein zweites Numpy-Array, das 200 mal die Zahl 0 enthält und speichern Sie es in der Variable `my_zeroes`. Erstellen Sie ein drittes Numpy-Array, das 300 mal die Zahl 5 enthält und speichern Sie es in der Variable `my_fives`. Fügen Sie die drei Arrays zusammen und speichern Sie das Ergebnis in der Variable `my_big_array`.

**Hinweis:** Sie können die Funktion `np.concatenate()` verwenden, um Numpy-Arrays zusammenzuführen.

```python
arr1 = np.array([1, 2, 3])
arr2 = np.array([4, 5, 6])
big_array = np.concatenate([arr1, arr2])
print(big_array) # Gibt [1, 2, 3, 4, 5, 6] aus
```

Berechnen Sie den Mittelwert des Arrays `my_big_array` und speichern Sie das Ergebnis in der Variable `my_big_array_mean`.

**Hinweis:** Sie können die Funktion `np.mean()` verwenden, um den Mittelwert eines Numpy-Arrays zu berechnen.

```python
arr = np.array([1, 2, 3])
mean = np.mean(arr)
print(mean) # Gibt 2.0 aus
```


```python
my_ones = np.ones(...)
my_zeroes = np.zeros(...)
my_fives = np.full(..., ...)
my_big_array = ...
my_big_array_mean = ...
```

```python
grader.check("q5")
```

## Funktionen

NumPy ist eine Bibliothek für numerische Berechnungen in Python. Sie bietet eine Vielzahl von Funktionen und Methoden, die das Arbeiten mit Arrays und Matrizen erleichtern. Hier werden wir einige der grundlegenden Funktionen und Methoden von NumPy kennenlernen.

- np.array(): Erstellt ein Numpy-Array aus einer Liste oder einem Tupel (has haben wir bereits gesehen)
- np.mean(): Berechnet den Mittelwert eines Numpy-Arrays
- np.unique(): Gibt die eindeutigen Elemente eines Numpy-Arrays zurück
- np.concatenate(): Fügt Numpy-Arrays zusammen
- np.arange(): Erstellt ein Numpy-Array mit einer Sequenz von Zahlen
- np.zeros(): Erstellt ein Numpy-Array mit Nullen
- np.ones(): Erstellt ein Numpy-Array mit Einsen
- np.log(): Berechnet den natürlichen Logarithmus eines Numpy-Arrays
- np.exp(): Berechnet die Exponentialfunktion eines Numpy-Arrays
- np.sum(): Berechnet die Summe eines Numpy-Arrays
- np.prod(): Berechnet das Produkt eines Numpy-Arrays
- np.min(): Gibt das Minimum eines Numpy-Arrays zurück
- np.max(): Gibt das Maximum eines Numpy-Arrays zurück
- np.argmax(): Gibt den Index des Maximums eines Numpy-Arrays zurück
- np.argmin(): Gibt den Index des Minimums eines Numpy-Arrays zurück
- np.sort(): Sortiert ein Numpy-Array
- np.quantile(): Berechnet das Quantil eines Numpy-Arrays
- np.median(): Berechnet den Median eines Numpy-Arrays
- np.std(): Berechnet die Standardabweichung eines Numpy-Arrays


**Aufgabe 6:** In dieser Aufgaben möchten wir die Anzahl von infektiösen Krankheitsfällen in einem Monat in einer Stadt. Wie viele Personen binnen eines Monats infiziert werden ist ungewiss und kann nicht genau vorhergesagt werden. Wir haben jedoch unsere Erfahrung aus der Vergangenheit in der Form einer Wahrhscheinlichkeitsverteilung zusammengefasst. 

Die mögliche Anzahl von Krankheitsfällen ist in der Variable `cases` gespeichert und läuft von 0 bis 300. Die Log-Wahrscheinlichkeiten für die Anzahl von Krankheitsfällen sind in der Variable `log_probs` gespeichert.

- Berechnen Sie die erwartete Anzahl von Krankheitsfällen und speichern Sie das Ergebnis in der Variable `expected_cases`.
- Finden Sie die Anzahl von Krankheitsfällen, für die die Wahrscheinlichkeit am größten ist und speichern Sie das Ergebnis in der Variable `most_likely_cases`.
- Berechnen Sie die Wahrscheinlichkeit, daß die Anzahl von Krankheitsfällen kleiner oder gleich 100 ist und speichern Sie das Ergebnis in der Variable `prob_less_than_100`.

Der Erwartungswert einer Zufallsvariablen ist definiert als die Summe der Produkte der Werte der Zufallsvariablen und ihrer Wahrscheinlichkeiten. Für eine diskrete Zufallsvariable X mit den Werten $x_1, x_2, \ldots , x_k$ und den Wahrscheinlichkeiten $p1, p2, \ldots, p_k$ ist der Erwartungswert E(X) gegeben durch:

$$
E(X) = x_1 p_1 + x_2  p_2 + \ldots + x_k p_k
$$

```python
from scipy.stats import poisson

cases = np.arange(0, 301)
log_probs = poisson.logpmf(cases, 100)
probs = ...
expected_cases = np.sum(...)
most_likely_cases = cases[...]
prob_less_than_100 = np.sum(...)
```

```python
grader.check("q6")
```

Tests for q1 (points: None)
```
>>> np.alltrue(a == np.array([1, 3, 8]))
True
```
```
>>> np.alltrue(b == np.array([3, 5, 10]))
True
```


Tests for q2 (points: None)
```
>>> np.alltrue(np.isclose(x1_div_x2, np.array([0.5, 0.6, 2.66666667, 0.85714286])))
True
```


Tests for q3 (points: None)
```
>>> third_element == 'c'
True
```


Tests for q4 (points: None)
```
>>> len(sent_arr) == 737
True
```
```
>>> sent_arr[:10] == np.array(['T', 'h', 'e', 'r', 'e', ' ', 'w', 'a', 's', ' '])
array([False, False, False, False, False, False, False, False, False,
       False])
```
```
>>> num_unique_chars == 33
True
```


Tests for q5 (points: None)
```
>>> len(my_ones) == 500
True
```
```
>>> np.alltrue(my_ones == np.ones(500))
True
```
```
>>> len(my_zeroes) == 200
True
```
```
>>> np.alltrue(my_zeroes == np.zeros(200))
True
```
```
>>> len(my_fives) == 100
True
```
```
>>> np.alltrue(my_fives == np.full(100, 5))
True
```
```
>>> len(my_big_array) == 800
True
```
```
>>> np.alltrue(my_big_array == np.concatenate([np.ones(500), np.zeros(200), np.full(100, 5)]))
True
```
```
>>> np.isclose(my_big_array_mean, 1.25)
True
```


Tests for q6 (points: None)
```
>>> len(probs) == 301
True
```
```
>>> np.isclose(np.sum(probs), 1)
True
```
```
>>> np.isclose(expected_cases, 100)
True
```
```
>>> np.isclose(most_likely_cases, 99)
True
```
```
>>> np.isclose(prob_less_than_100, 0.5265621985300244)
True
```
